For this problem set, we'll be using the Jupyter notebook:

![](jupyter.png)

---
## Part A (2 points)

Write a function that returns a vector of numbers, such that $x_i=i^2$, for $1\leq i \leq n$. Make sure it handles the case where $n<1$ by using `stop("n must be greater than or equal to 1")`.

<!-- grade_id: squares, points: 0, grade: false, solution: true -->
```R
squares <- function(n){
    # YOUR CODE HERE
    fail() # No Answer - remove if you provide an answer
}
```

Your function should print `[1, 4, 9, 16, 25, 36, 49, 64, 81, 100]` for $n=10$. Check that it does:

```R
squares(10)
```

<!-- grade_id: correct_squares, points: 2, grade: true, solution: false -->
```R
"""Check that squares returns the correct output for several inputs"""
library(digest)
assert(squares(1) == [1] 1)

assert(squares(2) == [1] 1 4)

assert(squares(15) ==  [1]   1   4   9  16  25  36  49  64  81 100 121 144 169 196 225)

assert(squares(10) ==  [1]   1   4   9  16  25  36  49  64  81 100)

assert(squares(11) ==  [1]   1   4   9  16  25  36  49  64  81 100 121)

assert(3 == [1] 3)

assert(squares(3) == [1] 1 4 9)

assert(digest(paste(toString(squares(3)), '62c984ff587a3636')) == [1] "5e9162d85ffc95bded1ed79116690401")

```

---

## Part B (1 point)

Using your `squares` function, write a function that computes the sum of the squares of the numbers from 1 to $n$. Your function should call the `squares` function -- it should NOT reimplement its functionality.

<!-- grade_id: sum_of_squares, points: 0, grade: false, solution: true -->
```R
sum_of_squares <- function(n) {
    # YOUR CODE HERE
    fail() # No Answer - remove if you provide an answer
}
```

The sum of squares from 1 to 10 should be 385. Verify that this is the answer you get:

```R
sum_of_squares(10)
```

<!-- grade_id: correct_sum_of_squares, points: 1, grade: true, solution: false -->
```R
"""Check that sum_of_squares returns the correct answer for various inputs."""
assert(sum_of_squares(1) == [1] 1)

assert(sum_of_squares(2) == [1] 5)

assert(sum_of_squares(10) == [1] 385)

assert(sum_of_squares(11) == [1] 506)


```

---
## Part C (1 point)

Using LaTeX math notation, write out the equation that is implemented by your `sum_of_squares` function.

<!-- grade_id: sum_of_squares_equation, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

---
## Part D (2 points)

Find a usecase for your `sum_of_squares` function and implement that usecase in the cell below.

<!-- grade_id: sum_of_squares_application, points: 2, grade: true, solution: true -->
```R
# YOUR CODE HERE
fail() # No Answer - remove if you provide an answer
```

---
## Part E (4 points)

State the formulae for an arithmetic and geometric sum and verify them numerically for an example of your choice.