For this problem set, we'll be using the Jupyter notebook:

![](jupyter.png)

---
## Part A (2 points)

Write a function that returns a vector of numbers, such that $x_i=i^2$, for $1\leq i \leq n$. Make sure it handles the case where $n<1$ by using `stop("n must be greater than or equal to 1")`.

<!-- grade_id: squares, points: 0, grade: false, solution: true -->
```R
squares <- function(n){
    # YOUR CODE HERE
    fail() # No Answer - remove if you provide an answer
}
```

Your function should print `[1, 4, 9, 16, 25, 36, 49, 64, 81, 100]` for $n=10$. Check that it does:

```R
squares(10)
```

<!-- grade_id: correct_squares, points: 2, grade: true, solution: false -->
```R
library(digest)
"assert(squares(1) == 1)
 assert(nchar(squares(1)) == 1)
 print('Success!')
"

"assert(squares(2) == 1 4)
 assert(nchar(squares(2)) == 1 1)
 print('Success!')
"

"assert(squares(15) ==  1   4   9  16  25  36  49  64  81 100 121 144 169 196 225)
 assert(nchar(squares(15)) ==  1 1 1 2 2 2 2 2 2 3 3 3 3 3 3)
 print('Success!')
"

"assert(squares(10) ==  1   4   9  16  25  36  49  64  81 100)
 assert(nchar(squares(10)) ==  1 1 1 2 2 2 2 2 2 3)
 print('Success!')
"

"assert(squares(11) ==  1   4   9  16  25  36  49  64  81 100 121)
 assert(nchar(squares(11)) ==  1 1 1 2 2 2 2 2 2 3 3)
 print('Success!')
"

"assert(3 == 3)
 assert(nchar(3) == 1)
 print('Success!')
"

"assert(squares(3) == 1 4 9)
 assert(nchar(squares(3)) == 1 1 1)
 print('Success!')
"

"assert(digest(paste(squares(3), '9d8ad6edce7137bf')) == 4a7ee3acf082614a5d5c81dcfbb0009e)
 assert(digest(paste(nchar(squares(3)), '9d8ad6edce7137bf')) == 4b3da6b75283e468ae94d798a3490aaa)
 print('Success!')
"

```

---

## Part B (1 point)

Using your `squares` function, write a function that computes the sum of the squares of the numbers from 1 to $n$. Your function should call the `squares` function -- it should NOT reimplement its functionality.

<!-- grade_id: sum_of_squares, points: 0, grade: false, solution: true -->
```R
sum_of_squares <- function(n) {
    # YOUR CODE HERE
    fail() # No Answer - remove if you provide an answer
}
```

The sum of squares from 1 to 10 should be 385. Verify that this is the answer you get:

```R
sum_of_squares(10)
```

<!-- grade_id: correct_sum_of_squares, points: 1, grade: true, solution: false -->
```R
"assert(sum_of_squares(1) == 1)
 assert(nchar(sum_of_squares(1)) == 1)
 print('Success!')
"

"assert(sum_of_squares(2) == 5)
 assert(nchar(sum_of_squares(2)) == 1)
 print('Success!')
"

"assert(sum_of_squares(10) == 385)
 assert(nchar(sum_of_squares(10)) == 3)
 print('Success!')
"

"assert(sum_of_squares(11) == 506)
 assert(nchar(sum_of_squares(11)) == 3)
 print('Success!')
"


```

---
## Part C (1 point)

Using LaTeX math notation, write out the equation that is implemented by your `sum_of_squares` function.

<!-- grade_id: sum_of_squares_equation, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

---
## Part D (2 points)

Find a usecase for your `sum_of_squares` function and implement that usecase in the cell below.

<!-- grade_id: sum_of_squares_application, points: 2, grade: true, solution: true -->
```R
# YOUR CODE HERE
fail() # No Answer - remove if you provide an answer
```

---
## Part E (4 points)

State the formulae for an arithmetic and geometric sum and verify them numerically for an example of your choice.