For this problem set, we'll be using the Jupyter notebook:

![](jupyter.png)

---
## Part A (2 points)

Write a function that returns a list of numbers, such that $x_i=i^2$, for $1\leq i \leq n$. Make sure it handles the case where $n<1$ by raising a `ValueError`.

<!-- grade_id: squares, points: 0, grade: false, solution: true -->
```python
def squares(n):
    """Compute the squares of numbers from 1 to n, such that the 
    ith element of the returned list equals i^2.
    
    """
    # YOUR CODE HERE
    if n < 1:
        raise
    return [i**2 for i in range(1, n+1)]
```

Your function should print `[1, 4, 9, 16, 25, 36, 49, 64, 81, 100]` for $n=10$. Check that it does:

```python
squares(10)
```

Output:
```
[1, 4, 9, 16, 25, 36, 49, 64, 81, 100]
```

<!-- grade_id: correct_squares, points: 1, grade: true, solution: false -->
```python
"""Check that squares returns the correct output for several inputs"""
from hashlib import sha1
assert sha1(str(squares(1)).encode('utf-8')+b'c4e7f2524fb9f618').hexdigest() == '0e60ef63a988f38f6e7f65508a9ab3b4c1e696ef'

assert sha1(str(squares(2)).encode('utf-8')+b'09adf5a36d10e777').hexdigest() == '51d07cf66e4c8d7268010258bd453387b43b3947'

assert sha1(str(squares(15)).encode('utf-8')+b'1d08dd3db62403bb').hexdigest() == '4a9153fa5969d25469a4535e09b42b7a473ce5a8'

assert sha1(str(squares(10)).encode('utf-8')+b'85dd8c35c4e94fda').hexdigest() == '55904776dabd36ba56f28112ae26b433bd3fd124'

assert sha1(str(squares(11)).encode('utf-8')+b'6fec31bb183c714d').hexdigest() == 'eaaca6b1df8bfefffa1573d2113d795896dc3884'

assert sha1(str(3).encode('utf-8')+b'b56cf517b0647c1e').hexdigest() == 'f5678997d5b43a20cb71380a01fe2723b42910af'

```

<!-- grade_id: squares_invalid_input, points: 1, grade: true, solution: false -->
```python
"""Check that squares raises an error for invalid inputs"""
try:
    squares(0)
except ValueError:
    pass
else:
    raise AssertionError("did not raise")

try:
    squares(-4)
except ValueError:
    pass
else:
    raise AssertionError("did not raise")
```

---

## Part B (1 point)

Using your `squares` function, write a function that computes the sum of the squares of the numbers from 1 to $n$. Your function should call the `squares` function -- it should NOT reimplement its functionality.

<!-- grade_id: sum_of_squares, points: 0, grade: false, solution: true -->
```python
def sum_of_squares(n):
    """Compute the sum of the squares of numbers from 1 to n."""
    # YOUR CODE HERE
    return sum(squares(n))
```

The sum of squares from 1 to 10 should be 385. Verify that this is the answer you get:

```python
sum_of_squares(10)
```

Output:
```
385
```

<!-- grade_id: correct_sum_of_squares, points: 0.5, grade: true, solution: false -->
```python
"""Check that sum_of_squares returns the correct answer for various inputs."""
assert sha1(str(sum_of_squares(1)).encode('utf-8')+b'49ce9d4a4aae97ae').hexdigest() == '110cba0cfa7f312eac7426e5c3116ffa0d4d3011'

assert sha1(str(sum_of_squares(2)).encode('utf-8')+b'bf5fcd3f1fd7db49').hexdigest() == '31b0b1beb3b8f0a041993e1d6d2e32d458151e1c'

assert sha1(str(sum_of_squares(10)).encode('utf-8')+b'84316344c29641f3').hexdigest() == 'ec68bacd9554581aa7182e3dc0889f13814bf1c2'

assert sha1(str(sum_of_squares(11)).encode('utf-8')+b'aa7134c4120bbdd2').hexdigest() == 'b39245a236afe5fe923f7ffc727d84ccd7a4567c'


```

<!-- grade_id: sum_of_squares_uses_squares, points: 0.5, grade: true, solution: false -->
```python
"""Check that sum_of_squares relies on squares."""
orig_squares = squares
del squares
try:
    sum_of_squares(1)
except NameError:
    pass
else:
    raise AssertionError("sum_of_squares does not use squares")
finally:
    squares = orig_squares
```

---
## Part C (1 point)

Using LaTeX math notation, write out the equation that is implemented by your `sum_of_squares` function.

<!-- grade_id: sum_of_squares_equation, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

---
## Part D (2 points)

Find a usecase for your `sum_of_squares` function and implement that usecase in the cell below.

<!-- grade_id: sum_of_squares_application, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

---
## Part E (4 points)

State the formulae for an arithmetic and geometric sum and verify them numerically for an example of your choice.