For this problem set, we'll be using the Jupyter notebook:

![](jupyter.png)

---
## Part A (2 points)

Write a function that returns a list of numbers, such that $x_i=i^2$, for $1\leq i \leq n$. Make sure it handles the case where $n<1$ by raising a `ValueError`.

<!-- grade_id: squares, points: 0, grade: false, solution: true -->
```python
def squares(n):
    """Compute the squares of numbers from 1 to n, such that the 
    ith element of the returned list equals i^2.
    
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

Your function should print `[1, 4, 9, 16, 25, 36, 49, 64, 81, 100]` for $n=10$. Check that it does:

```python
squares(10)
```

<!-- grade_id: correct_squares, points: 1, grade: true, solution: false -->
```python
"""Check that squares returns the correct output for several inputs"""
from hashlib import sha1
assert str(squares(1)) == '[1]', 'Value not correct.'

assert str(squares(2)) == '[1, 4]', 'Value not correct.'

assert str(squares(15)) == '[1, 4, 9, 16, 25, 36, 49, 64, 81, 100, 121, 144, 169, 196, 225]', 'Value not correct.'

assert str(squares(10)) == '[1, 4, 9, 16, 25, 36, 49, 64, 81, 100]', 'Value not correct.'

assert str(squares(11)) == '[1, 4, 9, 16, 25, 36, 49, 64, 81, 100, 121]', 'Value not correct.'

assert str(3) == '3', 'Value not correct.'

assert str(squares(3)) == '[1, 4, 9]', 'Value not correct.'

assert sha1(str(squares(3)).encode("utf-8")+b"96db95215791ce30").hexdigest() == '2ce8fc74da9384ec42a42da231dca4cf9b7e54b6', 'Value not correct.'

```

<!-- grade_id: squares_invalid_input, points: 1, grade: true, solution: false -->
```python
"""Check that squares raises an error for invalid inputs"""
def test_func_throws(func, ErrorType):
    try:
        func()
    except ErrorType:
        return True
    else:
        print('Did not raise right type of error!')
        return False
    
from hashlib import sha1
assert str(test_func_throws(lambda : squares(0), ValueError)) == 'True', 'Value not correct.'

assert str(test_func_throws(lambda : squares(-4), ValueError)) == 'True', 'Value not correct.'


```

---

## Part B (1 point)

Using your `squares` function, write a function that computes the sum of the squares of the numbers from 1 to $n$. Your function should call the `squares` function -- it should NOT reimplement its functionality.

<!-- grade_id: sum_of_squares, points: 0, grade: false, solution: true -->
```python
def sum_of_squares(n):
    """Compute the sum of the squares of numbers from 1 to n."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

The sum of squares from 1 to 10 should be 385. Verify that this is the answer you get:

```python
sum_of_squares(10)
```

<!-- grade_id: correct_sum_of_squares, points: 0.5, grade: true, solution: false -->
```python
"""Check that sum_of_squares returns the correct answer for various inputs."""
from hashlib import sha1
assert str(sum_of_squares(1)) == '1', 'Value not correct.'

assert str(sum_of_squares(2)) == '5', 'Value not correct.'

assert str(sum_of_squares(10)) == '385', 'Value not correct.'

assert str(sum_of_squares(11)) == '506', 'Value not correct.'


```

<!-- grade_id: sum_of_squares_uses_squares, points: 0.5, grade: true, solution: false -->
```python
"""Check that sum_of_squares relies on squares."""

orig_squares = squares
del squares

from hashlib import sha1
assert str(test_func_throws(lambda : sum_of_squares(1), NameError)) == 'False', 'Value not correct.'


squares = orig_squares

```

---
## Part C (1 point)

Using LaTeX math notation, write out the equation that is implemented by your `sum_of_squares` function.

<!-- grade_id: sum_of_squares_equation, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

---
## Part D (2 points)

Find a usecase for your `sum_of_squares` function and implement that usecase in the cell below.

<!-- grade_id: sum_of_squares_application, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

---
## Part E (4 points)

State the formulae for an arithmetic and geometric sum and verify them numerically for an example of your choice.